```python
# Initialize Otter
import otter
grader = otter.Notebook("lab4_clustering_coefficient_er.ipynb")
```

```python
#!pip install --upgrade networkx
```

```python
from IPython.core.display import HTML
from datascience import *

import matplotlib
matplotlib.use('Agg')
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
import os
plt.style.use('fivethirtyeight')

import networkx as nx

import warnings
warnings.filterwarnings('ignore')
```

# Lab4: Clustering coefficient and the Erdos-Renyi random graph Model

In this lab, we are going to test a hypothesis about clustering in social networks. 

From the lectures, we learned that social networks tend to be clustered because of mechanisms such as triadic closure. However, how can we be sure that clustering really exists in real-life networks? It seems possible that the structure of any network is a result of random processes rather than clustering. We need to perform some hypothesis testing to examine this possibility.

First, similar to what we did in Lab 3, we will will set up a null model and an alternative model. Our alternative model will be a real-life social network, and our null model will be random generated networks with the same number of nodes and the same level of connectivity (i.e. the same average degree). Then, we can compare the two networks to see if the real-life social network is more clustered than the random networks. 

We will use one of the 84 Add Health networks for the alternative model, and 1000 Erdos-Renyi(ER) networks that we learned in lecture for the null model. Clustering will be measured by the clustering coefficient, defined as the probability that two randomly selected friends of the node are friends with each other (You can chcek out this [demo](https://datahub.berkeley.edu/user/xinghe_pan/lab/workspaces/auto-K/tree/demog180-fa2024/lecture/20240905_clusteringcoef/clustering_coefficient.ipynb) on clustering coefficient).

Our hypotheses are: 
    
**Null hypothesis ($H_0$): the level of clustering is the same in the Add Health network and the random networks.**  

**Alternative hypothesis ($H_1$): the level of clustering is higher in the Add Health network than in the random networks.**

### Add Health Networks

Let's start by reading in one arbitrary Add Health network:

```python
def read_add_health_network(network_id):
    # this file was downloaded from
    # http://moreno.ss.uci.edu/data.html#adhealth
    edge_file = os.path.join("data", "comm" + str(network_id) + ".dat")
    with open(edge_file, 'r') as f:
        edge_lines = f.readlines()
        
    network = nx.parse_edgelist(edge_lines, nodetype=int, data=[('activity_level', float)])
    
    # note that we call the to_undirected method to ensure we get an undirected network
    return(network.to_undirected())
```

```python
# read in network #10
add_health_network = read_add_health_network(9)
```

## Question 1:

Calculate network statistics, including number of nodes, average degree, and clustering coefficient for `add_health_network`.

```python
num_nodes = ...
avg_degree = ...
clustering_coef = ...
```

```python
grader.check("q1")
```

The clustering coefficient of this Add Health network is:

```python
clustering_coef
```

## ER Networks

We learned that the ER model has two parameters: 

* $n$ - the number of nodes in the network
* $p$ - for each pair of nodes, the probability that those nodes are connected by an edge

## Question 2: 

Generate an ER model using the parameters $n = 10$ and $p = 0.3$:

```python

first_er = nx.erdos_renyi_graph(...) 
nx.draw(first_er)
```

```python
grader.check("q2")
```

The ER model can also be parametrized by the average degree, or $\bar{d}$, because average degree is directly related to $n$ and $p$:

$$\bar{d} = p (n-1)$$

that is, 

$$p = \frac{\bar{d}}{n-1}$$


## Question 3:

Calculate the average degree of an ER model with parameters $n=10$ and $p=0.3$:

```python
n_q3 = ...
p_q3 = ...
dbar_q3 = ...
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

## Question 4:

Let's write a custom function to generate ER networks using parameters $n$ and $\bar{d}$ for our purpose:

```python
def gen_er_by_avg_degree(n, dbar):
    return(nx.erdos_renyi_graph(...)) 
```

<!-- END QUESTION -->

## Question 5

Now, we can generate 1000 ER networks, using the number of nodes and the average degree of the Add Health network as values of $n$ and $\bar{d}$, respectively. Then, calculate the clustering coefficient for each network and store in an array:

```python
n = ...
dbar = ...

cc_er = make_array()

for _ in range(1000):
    er = gen_er_by_avg_degree(...)
    cc_er = np.append(cc_er, ...)

```

```python
grader.check("q5")
```

Let's check the distribution of clustering coefficients in the ER models:

```python
Table().with_column('cc_er', cc_er).hist()
```

It looks like the majority of ER networks has clustering coefficients below 0.01, which seems much smaller than that of the Add Health network we selected. Let's calculate an empirical $p$ value to be sure:

```python
emp_p_value = np.mean(cc_er >= clustering_coef)
emp_p_value
```

<!-- BEGIN QUESTION -->

## Question 6

Based on the empirical $p$ value above, what can we conclude about our null and alternative hypotheses based on a significance level of 0.05?  

_Type your answer here, replacing this text._

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> np.isclose(num_nodes, 248.0)
True
```
```
>>> np.isclose(clustering_coef, 0.313195)
True
```


Tests for q2 (points: None)
```
>>> first_er.number_of_nodes() == 10
True
```


Tests for q3 (points: None)
```
>>> np.isclose(dbar_q3, 2.69999)
True
```


Tests for q5 (points: None)
```
>>> len(cc_er) == 1000
True
```
```
>>> np.isclose(n, 248)
True
```
```
>>> np.isclose(dbar, 8.096774)
True
```
